
::::{grid} 1 1 2 2

:::{card}
:header: **Lycée Français de Tananarive**
```{figure} https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px-2.png
:name: myFigure
:width: 300px
:align: center
```
:::

:::{card}
:header: **Numérique et Sciences Informatiques - Serveur Jupyter**
```{figure} https://galilee.ent.auvergnerhonealpes.fr/lectureFichiergw.do?ID_FICHIER=3782
:width: 200px
:align: center
```
:::

::::

# NoteBook : Epreuve pratique n°3 
## Exercice 2

:::{warning} Liste chainée

On veut écrire une classe pour gérer une file à l’aide d’une liste chaînée. On dispose d’une classe Maillon permettant la création d’un maillon de la chaîne, celui-ci étant constitué d’une valeur et d’une référence au maillon suivant de la chaîne :

```python
class Maillon :
    def __init__(self,v) :
        self.valeur = v
        self.suivant = None
```

:::{tip} Compléter la classe
Compléter la classe File suivante où l’attribut dernier_file contient le maillon correspondant à l’élément arrivé en dernier dans la file.

Vous compléterez les méthodes :

- `enfile(self,element)` : Opérateur, Ajoute un élément à la file
- `est_vide(self)` : Prédicat : Teste si la file est vide
- `defile(self)`: Opérateur, défile un élément et renvoie sa valeur

<!-- grade_id: cell-a7b0d1776b7b5a41, points: 0, grade: false, solution: true -->
```python
class Maillon :
    def __init__(self,v) :
        self.valeur = v
        self.suivant = None

class File :
    def __init__(self) :
        self.dernier_file = None

    def enfile(self,element) :
        nouveau_maillon = Maillon(element)
        nouveau_maillon.suivant = self.dernier_file
        self.dernier_file = nouveau_maillon

    def est_vide(self) :
        if self.dernier_file == None:
            return True
        return False

    def affiche(self) :
        maillon = self.dernier_file
        affichage = ""
        while maillon != None :
            affichage = affichage + str(maillon.valeur)
            maillon = maillon.suivant
        return affichage

    def defile(self) :
        if not self.est_vide() :
            if self.dernier_file.suivant == None :
                resultat = self.dernier_file.valeur
                self.dernier_file = None
                return resultat
            maillon = self.dernier_file
            while maillon.suivant.suivant != None :
                maillon = maillon.suivant
            resultat = maillon.suivant.valeur
            maillon.suivant = None
            return resultat
        return None

F = File()
F.enfile(4)
F.enfile(5)
F.defile()
F.enfile(1)
F.enfile(2)
F.enfile(3)
F.defile()
```

Output:
```
5
```

<!-- grade_id: cell-350e40a40600e092, points: 2, grade: true, solution: false -->
```python
F = File()
F.enfile(3)
assert F.affiche() == "3"
F.enfile(2)
assert F.affiche() == "23"
F.enfile(6)
assert F.affiche() == "623"
```

<!-- grade_id: cell-97126dfd5a5e4835, points: 1, grade: true, solution: false -->
```python
F = File()
assert F.est_vide()
F.enfile(3)
assert not F.est_vide()
```

<!-- grade_id: cell-e9240580b6f439bf, points: 2, grade: true, solution: false -->
```python
F = File()
F.enfile(4)
F.enfile(5)
assert F.defile() == 4
F.enfile(1)
F.enfile(2)
F.enfile(3)
assert F.defile() == 5
assert F.defile() == 1
```

:::{card}
:header: **Equipe Pédagogique LFT-NSI**
ressource éducative libre distribuée sous [Licence Creative Commons Attribution](http://creativecommons.org/licenses/by-nc-sa/4.0/)

```{figure} https://licensebuttons.net/l/by-nc-sa/4.0/88x31.png
:alt: Beautiful Sunset
:width: 80px
:align: center
```
:::