
::::{grid} 1 1 2 2

:::{card}
:header: **Lycée Français de Tananarive**
```{figure} https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px-2.png
:name: myFigure
:width: 300px
:align: center
```
:::

:::{card}
:header: **Numérique et Sciences Informatiques - Serveur Jupyter**
```{figure} https://galilee.ent.auvergnerhonealpes.fr/lectureFichiergw.do?ID_FICHIER=3782
:width: 200px
:align: center
```
:::

::::

# NoteBook : Epreuve pratique n°3 
## Exercice 1

:::{warning} Un arbre binaire
Dans cet exercice, un arbre binaire de caractères est stocké sous la forme d’un dictionnaire où les clefs sont les caractères des nœuds de l’arbre et les valeurs, pour chaque clef, la liste des caractères des fils gauche et droit du nœud.

Par exemple, l’arbre est stocké dans

`a= {'F':['B','G'], 'B':['A','D'], 'A':['',''], 'D':['C','E'], 'C':['',''], 'E':['',''], 'G':['','I'], 'I':['','H'],'H':['','']}`

**Écrire une fonction récursive taille prenant en paramètres un arbre binaire arbre sous la forme d’un dictionnaire et un caractère lettre qui est la valeur du sommet de l’arbre, et qui renvoie la taille de l’arbre à savoir le nombre total de nœud.**

$\heartsuit$ : La lettre passée en paramètre appartient forcément à l'arbre, ce qui veut dire que l'arbre ne peut être vide

$\heartsuit$ : Un arbre réduit à sa racine aura donc une taille de 1.

On pourra distinguer les 4 cas où les deux « fils » du nœud sont '', le fils gauche seulement est '', le fils droit seulement est '', aucun des deux fils n’est ''.

Exemple :
```Python
taille(a, ’F’)
9
```

<!-- grade_id: cell-b00f3a3733649b5f, points: 0, grade: false, solution: true -->
```python
def taille(arbre,lettre):
    if arbre[lettre] == ["", ""]:
        return 1
    else:
        if not arbre[lettre][0] and not arbre[lettre][1]:
            return 1
        elif arbre[lettre][0] and not arbre[lettre][1]:
            return 1 + taille(arbre, arbre[lettre][0])
        elif not arbre[lettre][0] and arbre[lettre][1]:
            return 1 + taille(arbre, arbre[lettre][1])
        elif arbre[lettre][0] and arbre[lettre][1]:
            return 1 + taille(arbre, arbre[lettre][0]) + taille(arbre, arbre[lettre][1])
a= {'F':['B','G'], 'B':['A','D'], 'A':['',''], 'D':['C','E'], 'C':['',''], 'E':['',''], 'G':['','I'], 'I':['','H'],'H':['','']}
taille (a,"F" )
```

Output:
```
9
```

<!-- grade_id: cell-98c481cc6bc1a280, points: 1, grade: true, solution: false -->
```python
a1 = {'A' : ['','']}
assert taille(a1, "A") == 1
```

<!-- grade_id: cell-0320cd714fa3f0e1, points: 1, grade: true, solution: false -->
```python
a2 = {'A': ['B', ''], 'B' : ['','']}
assert taille(a2,'A') == 2
assert taille (a2,'B') == 1
```

<!-- grade_id: cell-72135151bb68b78f, points: 1, grade: true, solution: false -->
```python
a3 = {'A' : ['','B'] , 'B' : ['C','D'], 'C' : ['', ''], 'D' : ['', '']}
assert taille(a3, 'A') == 4
assert taille(a3, 'B') == 3
assert taille(a3, 'C') == 1
```

<!-- grade_id: cell-3892feb2ffd0e793, points: 2, grade: true, solution: false -->
```python
a= {'F':['B','G'], 'B':['A','D'], 'A':['',''], 'D':['C','E'], 'C':['',''], 'E':['',''], 'G':['','I'], 'I':['','H'],'H':['','']}
assert taille (a,"F" ) == 9
assert taille (a, 'D') == 3
```

:::{card}
:header: **Equipe Pédagogique LFT-NSI**
ressource éducative libre distribuée sous [Licence Creative Commons Attribution](http://creativecommons.org/licenses/by-nc-sa/4.0/)

```{figure} https://licensebuttons.net/l/by-nc-sa/4.0/88x31.png
:alt: Beautiful Sunset
:width: 80px
:align: center
```
:::